[1441826685] SERVICE ALERT: localhost;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441826737] SERVICE ALERT: mail.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441826859] SERVICE ALERT: ad.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441826911] SERVICE ALERT: application.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441826963] SERVICE ALERT: artifactory.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441827075] SERVICE ALERT: confluence.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441827128] SERVICE ALERT: git.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441827180] SERVICE ALERT: jenkins.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441827233] SERVICE ALERT: jira.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441856851] SERVICE ALERT: application.plank.life;CHECK SSH IPv4;CRITICAL;SOFT;1;Connection refused
[1441856971] SERVICE ALERT: application.plank.life;CHECK SSH IPv4;CRITICAL;SOFT;2;Connection refused
[1441857082] SERVICE ALERT: application.plank.life;SSH;CRITICAL;SOFT;1;Connection refused
[1441857091] SERVICE ALERT: application.plank.life;CHECK SSH IPv4;CRITICAL;HARD;3;Connection refused
[1441857201] SERVICE ALERT: application.plank.life;SSH;CRITICAL;SOFT;2;Connection refused
[1441857321] SERVICE ALERT: application.plank.life;SSH;CRITICAL;HARD;3;Connection refused
[1441858476] HOST ALERT: jira.plank.life;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.20)
[1441858484] HOST ALERT: jenkins.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441858489] HOST ALERT: confluence.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441858494] HOST ALERT: jira.plank.life;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.20)
[1441858495] HOST ALERT: confluence.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1441858499] HOST ALERT: git.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441858500] HOST ALERT: jenkins.plank.life;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1441858506] HOST ALERT: jenkins.plank.life;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1441858510] HOST ALERT: git.plank.life;UP;SOFT;2;PING OK - Packet loss = 80%, RTA = 0.55 ms
[1441858517] HOST ALERT: jira.plank.life;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1441859614] SERVICE ALERT: jira.plank.life;CHECK TIME;UNKNOWN;HARD;3;Connection refused
[1441859649] SERVICE ALERT: jira.plank.life;Check SSH_IPv4;UNKNOWN;HARD;3;Remote command execution failed: ssh: connect to host 10.0.0.20 port 22: Connection refused
[1441860224] SERVICE ALERT: jira.plank.life;CHECK TIME;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1441860415] HOST ALERT: jira.plank.life;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.20)
[1441860430] SERVICE ALERT: jira.plank.life;CHECK HOST ALIVE;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.20)
[1441860430] HOST ALERT: jira.plank.life;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.20)
[1441860452] HOST ALERT: jira.plank.life;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.0.0.20)
[1441860468] SERVICE ALERT: jira.plank.life;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (10.0.0.20)
[1441860471] HOST ALERT: jira.plank.life;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.0.0.20)
[1441860486] HOST ALERT: jira.plank.life;DOWN;HARD;5;CRITICAL - Host Unreachable (10.0.0.20)
[1441860528] SERVICE ALERT: jira.plank.life;PING;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (10.0.0.20)
[1441860550] SERVICE ALERT: jira.plank.life;CHECK HOST ALIVE;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (10.0.0.20)
[1441860588] SERVICE ALERT: jira.plank.life;PING;CRITICAL;SOFT;3;CRITICAL - Host Unreachable (10.0.0.20)
[1441860646] SERVICE ALERT: jira.plank.life;PING;CRITICAL;HARD;4;CRITICAL - Host Unreachable (10.0.0.20)
[1441860665] SERVICE ALERT: jenkins.plank.life;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 81%, RTA = 0.55 ms
[1441860670] SERVICE ALERT: jira.plank.life;CHECK HOST ALIVE;CRITICAL;HARD;3;CRITICAL - Host Unreachable (10.0.0.20)
[1441860678] HOST ALERT: git.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441860686] SERVICE ALERT: jenkins.plank.life;CHECK ROUTER ALIVE;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441860689] HOST ALERT: jenkins.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441860694] SERVICE ALERT: git.plank.life;CHECK SWITCH ALIVE;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441860695] HOST ALERT: git.plank.life;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.0.0.28)
[1441860703] HOST ALERT: jenkins.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1441860710] HOST ALERT: git.plank.life;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1441860710] HOST ALERT: jira.plank.life;UP;HARD;5;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1441860719] SERVICE ALERT: jenkins.plank.life;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1441860796] SERVICE ALERT: jenkins.plank.life;CHECK ROUTER ALIVE;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1441860804] SERVICE ALERT: git.plank.life;CHECK SWITCH ALIVE;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1441860859] SERVICE ALERT: jira.plank.life;Check SSH_IPv4;CRITICAL;HARD;3;CRITICAL - Plugin timed out after 10 seconds
[1441860949] SERVICE ALERT: jira.plank.life;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1441861267] SERVICE ALERT: jira.plank.life;CHECK HOST ALIVE;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1441861710] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441861761] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441861830] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441861874] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441861881] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441861925] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441861950] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441861979] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441861994] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862001] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862031] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441862045] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862083] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441862099] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862114] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862135] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441862151] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862165] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862203] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862204] HOST ALERT: jira.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441862210] HOST ALERT: jira.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.90 ms
[1441862219] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862255] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862258] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441862271] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862323] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862375] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862378] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441862498] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441862922] SERVICE ALERT: ad.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862930] SERVICE ALERT: application.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862939] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862947] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862956] SERVICE ALERT: git.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862964] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862973] SERVICE ALERT: jira.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862981] SERVICE ALERT: localhost;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441862998] SERVICE ALERT: mail.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863042] SERVICE ALERT: ad.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863050] SERVICE ALERT: application.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863059] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863067] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863076] SERVICE ALERT: git.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863084] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863093] SERVICE ALERT: jira.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863101] SERVICE ALERT: localhost;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863118] SERVICE ALERT: mail.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441863162] SERVICE ALERT: ad.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863170] SERVICE ALERT: application.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863179] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863187] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863196] SERVICE ALERT: git.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863204] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863213] SERVICE ALERT: jira.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863221] SERVICE ALERT: localhost;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863238] SERVICE ALERT: mail.plank.life;CHECK TOMCAT SESSIONS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441863892] SERVICE ALERT: ad.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863900] SERVICE ALERT: application.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863909] SERVICE ALERT: artifactory.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863925] SERVICE ALERT: confluence.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863934] SERVICE ALERT: git.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863942] SERVICE ALERT: jenkins.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863950] SERVICE ALERT: jira.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863967] SERVICE ALERT: localhost;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441863976] SERVICE ALERT: mail.plank.life;CHECK HEAP;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864012] SERVICE ALERT: ad.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864020] SERVICE ALERT: application.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864029] SERVICE ALERT: artifactory.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864045] SERVICE ALERT: confluence.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864054] SERVICE ALERT: git.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864062] SERVICE ALERT: jenkins.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864070] SERVICE ALERT: jira.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864087] SERVICE ALERT: localhost;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864096] SERVICE ALERT: mail.plank.life;CHECK HEAP;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864132] SERVICE ALERT: ad.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864140] SERVICE ALERT: application.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864149] SERVICE ALERT: artifactory.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864165] SERVICE ALERT: confluence.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864174] SERVICE ALERT: git.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864182] SERVICE ALERT: jenkins.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864190] SERVICE ALERT: jira.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864207] SERVICE ALERT: localhost;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864216] SERVICE ALERT: mail.plank.life;CHECK HEAP;UNKNOWN;HARD;3;(No output returned from plugin)
[1441864858] SERVICE ALERT: ad.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864863] HOST ALERT: confluence.plank.life;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.29)
[1441864866] SERVICE ALERT: application.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864874] SERVICE ALERT: artifactory.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864881] HOST ALERT: confluence.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1441864883] SERVICE ALERT: confluence.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864899] SERVICE ALERT: git.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864916] SERVICE ALERT: jenkins.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864924] SERVICE ALERT: jira.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864957] SERVICE ALERT: localhost;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864974] SERVICE ALERT: mail.plank.life;CHECK JVM;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441864978] SERVICE ALERT: ad.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864986] SERVICE ALERT: application.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441864994] SERVICE ALERT: artifactory.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865003] SERVICE ALERT: confluence.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865019] SERVICE ALERT: git.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865036] SERVICE ALERT: jenkins.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865044] SERVICE ALERT: jira.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865077] SERVICE ALERT: localhost;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865094] SERVICE ALERT: mail.plank.life;CHECK JVM;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865098] SERVICE ALERT: ad.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865106] SERVICE ALERT: application.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865114] SERVICE ALERT: artifactory.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865123] SERVICE ALERT: confluence.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865139] SERVICE ALERT: git.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865156] SERVICE ALERT: jenkins.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865164] SERVICE ALERT: jira.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865197] SERVICE ALERT: localhost;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865214] SERVICE ALERT: mail.plank.life;CHECK JVM;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865605] SERVICE ALERT: ad.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865613] SERVICE ALERT: application.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865621] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865629] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865637] SERVICE ALERT: git.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865662] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865695] SERVICE ALERT: jira.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865703] SERVICE ALERT: localhost;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865719] SERVICE ALERT: mail.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441865725] SERVICE ALERT: ad.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865733] SERVICE ALERT: application.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865741] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865749] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865757] SERVICE ALERT: git.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865782] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865815] SERVICE ALERT: jira.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865823] SERVICE ALERT: localhost;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865839] SERVICE ALERT: mail.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441865845] SERVICE ALERT: ad.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865853] SERVICE ALERT: application.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865861] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865869] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865877] SERVICE ALERT: git.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865902] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865935] SERVICE ALERT: jira.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865943] SERVICE ALERT: localhost;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441865959] SERVICE ALERT: mail.plank.life;CHECK TOMCAT MEMORY/THREADS;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866394] SERVICE ALERT: ad.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866402] SERVICE ALERT: application.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866410] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866418] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866426] SERVICE ALERT: git.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866434] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866442] SERVICE ALERT: jira.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866466] SERVICE ALERT: localhost;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866474] SERVICE ALERT: mail.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441866514] SERVICE ALERT: ad.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866522] SERVICE ALERT: application.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866530] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866538] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866546] SERVICE ALERT: git.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866554] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866562] SERVICE ALERT: jira.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866586] SERVICE ALERT: localhost;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866594] SERVICE ALERT: mail.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441866634] SERVICE ALERT: ad.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866642] SERVICE ALERT: application.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866650] SERVICE ALERT: artifactory.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866655] SERVICE ALERT: jira.plank.life;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1441866658] SERVICE ALERT: confluence.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866666] SERVICE ALERT: git.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866674] SERVICE ALERT: jenkins.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866682] SERVICE ALERT: jira.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866706] SERVICE ALERT: localhost;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441866709] SERVICE ALERT: jira.plank.life;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1441866714] SERVICE ALERT: mail.plank.life;CHECK TOMCAT APPLICATION;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867294] SERVICE ALERT: ad.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867310] SERVICE ALERT: application.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867318] SERVICE ALERT: artifactory.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867326] SERVICE ALERT: confluence.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867334] SERVICE ALERT: git.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867342] SERVICE ALERT: jenkins.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867350] SERVICE ALERT: jira.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867366] SERVICE ALERT: localhost;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867390] SERVICE ALERT: mail.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;1;(No output returned from plugin)
[1441867414] SERVICE ALERT: ad.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867430] SERVICE ALERT: application.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867438] SERVICE ALERT: artifactory.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867446] SERVICE ALERT: confluence.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867454] SERVICE ALERT: git.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867462] SERVICE ALERT: jenkins.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867470] SERVICE ALERT: jira.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867486] SERVICE ALERT: localhost;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867510] SERVICE ALERT: mail.plank.life;CHECK MYSQL HEALTH;UNKNOWN;SOFT;2;(No output returned from plugin)
[1441867534] SERVICE ALERT: ad.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867550] SERVICE ALERT: application.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867558] SERVICE ALERT: artifactory.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867566] SERVICE ALERT: confluence.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867574] SERVICE ALERT: git.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867582] SERVICE ALERT: jenkins.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867590] SERVICE ALERT: jira.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867606] SERVICE ALERT: localhost;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441867630] SERVICE ALERT: mail.plank.life;CHECK MYSQL HEALTH;UNKNOWN;HARD;3;(No output returned from plugin)
[1441869733] SERVICE ALERT: artifactory.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441869818] SERVICE ALERT: confluence.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441869870] SERVICE ALERT: git.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441869924] SERVICE ALERT: jenkins.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441869979] SERVICE ALERT: jira.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441870028] SERVICE ALERT: localhost;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441870081] SERVICE ALERT: mail.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441870228] SERVICE ALERT: ad.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441870279] SERVICE ALERT: application.plank.life;CHECK ALL DISKS;OK;HARD;3;DISK OK
[1441875959] SERVICE ALERT: jenkins.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876011] SERVICE ALERT: jira.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876063] SERVICE ALERT: localhost;CHECK DISK;OK;HARD;3;DISK OK
[1441876116] SERVICE ALERT: mail.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876261] SERVICE ALERT: ad.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876315] SERVICE ALERT: application.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876365] SERVICE ALERT: artifactory.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876453] SERVICE ALERT: confluence.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441876505] SERVICE ALERT: git.plank.life;CHECK DISK;OK;HARD;3;DISK OK
[1441878089] SERVICE ALERT: application.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878138] SERVICE ALERT: artifactory.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878228] SERVICE ALERT: confluence.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878275] SERVICE ALERT: git.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878334] SERVICE ALERT: jenkins.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878383] SERVICE ALERT: jira.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878438] SERVICE ALERT: localhost;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878491] SERVICE ALERT: mail.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878518] SERVICE ALERT: ad.plank.life;DHCP;OK;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878587] SERVICE ALERT: application.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878634] SERVICE ALERT: ad.plank.life;CHECK BREEZE;CRITICAL;HARD;3;Signal Strength at: 0%
[1441878648] SERVICE ALERT: artifactory.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878734] SERVICE ALERT: confluence.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878789] SERVICE ALERT: git.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878840] SERVICE ALERT: jenkins.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878891] SERVICE ALERT: jira.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878946] SERVICE ALERT: localhost;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441878999] SERVICE ALERT: mail.plank.life;DHCP;CRITICAL;HARD;3;Warning: This plugin must be either run as root or setuid root.
[1441879298] SERVICE ALERT: ad.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879349] SERVICE ALERT: application.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879401] SERVICE ALERT: artifactory.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879450] SERVICE ALERT: mail.plank.life;CHECK SPOP;OK;HARD;3;POP OK - 0.055 second response time on port 995 [+OK Dovecot (Ubuntu) ready.]
[1441879457] SERVICE ALERT: mail.plank.life;CHECK SPOP IPv4;OK;HARD;3;POP OK - 0.058 second response time on port 995 [+OK Dovecot (Ubuntu) ready.]
[1441879510] SERVICE ALERT: mail.plank.life;CHECK POP IPv4;OK;HARD;3;POP OK - 0.005 second response time on port 110 [+OK Dovecot (Ubuntu) ready.]
[1441879514] SERVICE ALERT: confluence.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879565] SERVICE ALERT: git.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879613] SERVICE ALERT: mail.plank.life;IMAP;OK;HARD;3;IMAP OK - 0.015 second response time on port 143 [* OK [CAPABILITY IMAP4rev1 LITERAL+ SASL-IR LOGIN-REFERRALS ID ENABLE IDLE STARTTLS LOGINDISABLED] Dovecot (Ubuntu) ready.]
[1441879619] SERVICE ALERT: jenkins.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879671] SERVICE ALERT: jira.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879723] SERVICE ALERT: localhost;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879775] SERVICE ALERT: mail.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441879870] SERVICE ALERT: mail.plank.life;CHECK IMAP IPv4;OK;HARD;3;IMAP OK - 0.014 second response time on port 143 [* OK [CAPABILITY IMAP4rev1 LITERAL+ SASL-IR LOGIN-REFERRALS ID ENABLE IDLE STARTTLS LOGINDISABLED] Dovecot (Ubuntu) ready.]
[1441881098] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881149] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881201] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881218] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881269] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881314] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881321] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881338] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881365] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881389] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881419] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881434] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881441] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881471] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881485] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881523] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881539] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881554] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881575] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441881591] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881605] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881643] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881659] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881695] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441881711] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881763] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441881815] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441882132] SERVICE ALERT: ad.plank.life;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882139] SERVICE ALERT: application.plank.life;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882149] SERVICE ALERT: artifactory.plank.life;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882163] SERVICE ALERT: confluence.plank.life;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882174] SERVICE ALERT: git.plank.life;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882182] SERVICE ALERT: jenkins.plank.life;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882199] SERVICE ALERT: jira.plank.life;CHECK HEAP;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1441882206] SERVICE ALERT: localhost;CHECK HEAP;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1441882226] SERVICE ALERT: mail.plank.life;CHECK HEAP;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1441889738] SERVICE ALERT: ad.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441889789] SERVICE ALERT: application.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441889841] SERVICE ALERT: artifactory.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441889954] SERVICE ALERT: confluence.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441890005] SERVICE ALERT: git.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441890059] SERVICE ALERT: jenkins.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441890111] SERVICE ALERT: jira.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441890163] SERVICE ALERT: localhost;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441890215] SERVICE ALERT: mail.plank.life;CHECK USERS;OK;HARD;3;USERS OK - 0 users currently logged in
[1441905814] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441905933] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441905937] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441905988] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441906040] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441906052] SERVICE ALERT: mail.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441906057] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441906108] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441906154] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441906160] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 1 users currently logged in
[1441906177] SERVICE ALERT: ad.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441906204] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 1 users currently logged in
[1441906228] SERVICE ALERT: application.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 1 users currently logged in
[1441906258] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 2 users currently logged in
[1441906274] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 2 users currently logged in
[1441906280] SERVICE ALERT: artifactory.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 2 users currently logged in
[1441906310] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 2 users currently logged in
[1441906324] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 2 users currently logged in
[1441906361] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;SOFT;1;USERS CRITICAL - 2 users currently logged in
[1441906378] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 2 users currently logged in
[1441906394] SERVICE ALERT: confluence.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 2 users currently logged in
[1441906430] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 2 users currently logged in
[1441906444] SERVICE ALERT: git.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 2 users currently logged in
[1441906481] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;SOFT;2;USERS CRITICAL - 2 users currently logged in
[1441906498] SERVICE ALERT: jenkins.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 2 users currently logged in
[1441906550] SERVICE ALERT: jira.plank.life;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 2 users currently logged in
[1441906601] SERVICE ALERT: localhost;CHECK USERS;CRITICAL;HARD;3;USERS CRITICAL - 2 users currently logged in
[1441906866] SERVICE ALERT: application.plank.life;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.88 ms