[1442220447] SERVICE ALERT: artifactory.plank.life;TCP;CRITICAL;SOFT;1;No route to host
[1442220448] HOST ALERT: mail.plank.life;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.4)
[1442220449] SERVICE ALERT: cisco;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 70%, RTA = 0.66 ms
[1442220457] HOST ALERT: confluence.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1442220468] HOST ALERT: confluence.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1442220468] SERVICE ALERT: git.plank.life;TIME;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1442220475] HOST ALERT: git.plank.life;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.0.0.28)
[1442220476] HOST ALERT: jira.plank.life;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1442220478] HOST ALERT: mail.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1442220488] HOST ALERT: jira.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1442220490] HOST ALERT: git.plank.life;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1001.07 ms
[1442220503] SERVICE ALERT: cisco;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1442220563] SERVICE ALERT: artifactory.plank.life;TCP;OK;SOFT;2;TCP OK - 0.001 second response time on port 5543
[1442221048] SERVICE ALERT: git.plank.life;TIME;UNKNOWN;HARD;3;Connection refused